# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 2-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with numbers and/or parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.  `'Switzerland17'` should be `'Switzerland'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
def answer_one():
    # YOUR CODE HERE
    data_energy = pd.read_excel("assets/Energy Indicators.xls", skiprows=17, skip_footer = 38)
    Energy = pd.DataFrame(data_energy)
    Energy = Energy[["Unnamed: 1", "Petajoules", "Gigajoules", "%"]]
    Energy.columns = ["Country", 'Energy Supply', 'Energy Supply per Capita', '% Renewable']
    Energy[["Country", 'Energy Supply', 'Energy Supply per Capita', '% Renewable']] = Energy[["Country", 'Energy Supply', 'Energy Supply per Capita', '% Renewable']].replace('...', np.NaN)
    Energy[['Country']] = Energy[['Country']].replace({"Republic of Korea": "South Korea",
                                                   "United States of America": "United States",
                                                   "United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
                                                   "China, Hong Kong Special Administrative Region": "Hong Kong", 
                                                       'Iran (Islamic Republic of)':'Iran'})
    Energy['Country'].str.replace(r"\(.*\)", "")
    Energy['Energy Supply'] = Energy['Energy Supply'] * 1000000
    
    data_gdp = pd.read_csv("assets/world_bank.csv", skiprows=4)
    gdp = pd.DataFrame(data_gdp)
    gdp[['Country Name']] = gdp[['Country Name']].replace({"Korea, Rep.": "South Korea", 
                                                 "Iran, Islamic Rep.": "Iran",
                                                 "Hong Kong SAR, China": "Hong Kong"})
    gdp = gdp[['Country Name', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    gdp.columns =  ['Country', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']
    
    data_scimen = pd.read_excel("assets/scimagojr-3.xlsx")
    ScimEn = pd.DataFrame(data_scimen)
    
    df_1 = pd.merge(ScimEn, Energy, how='inner',left_on='Country', right_on='Country')
    df_2 = pd.merge(df_1, gdp, how='inner', left_on='Country', right_on='Country')
    df_2 = df_2.set_index('Country')
    return df_2[:15]

print(answer_one())
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
China                  1     127050             126767     597237   
United States          2      96661              94747     792274   
Japan                  3      30504              30287     223024   
United Kingdom         4      20944              20357     206091   
Russian Federation     5      18534              18301      34266   
Canada                 6      17899              17620     215003   
Germany                7      17027              16831     140566   
India                  8      15005              14841     128763   
France                 9      13153              12973     130632   
South Korea           10      11983              11923     114675   
Italy                 11      10964              10794     111850   
Spain                 12       9428               9330     123336   
Iran                  13       8896               8819
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
def answer_two():
    # YOUR CODE HERE
    import pandas as pd
    import numpy as np

    Energy = pd.read_excel("assets/Energy Indicators.xls", skiprows=17, skip_footer = 38)
    Energy = Energy[["Unnamed: 1", "Petajoules", "Gigajoules", "%"]]
    Energy.columns = ["Country", 'Energy Supply', 'Energy Supply per Capita', '% Renewable']
    Energy[["Country", 'Energy Supply', 'Energy Supply per Capita', '% Renewable']] = Energy[["Country", 'Energy Supply', 'Energy Supply per Capita', '% Renewable']].replace('...', np.NaN)
    Energy[['Country']] = Energy[['Country']].replace({"Republic of Korea": "South Korea",
                                                   "United States of America": "United States",
                                                   "United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
                                                   "China, Hong Kong Special Administrative Region": "Hong Kong", 
                                                       'Iran (Islamic Republic of)':'Iran'})
    Energy['Country'].str.replace(r"\(.*\)", "")
    Energy['Energy Supply'] = Energy['Energy Supply'] * 1000000
    
    data_gdp = pd.read_csv("assets/world_bank.csv", skiprows=4)
    gdp = pd.DataFrame(data_gdp)
    gdp[['Country Name']] = gdp['Country Name'].replace({"Korea, Rep.": "South Korea", 
                                                 "Iran, Islamic Rep.": "Iran",
                                                 "Hong Kong SAR, China": "Hong Kong"})
    gdp = gdp[['Country Name', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    gdp.columns = ['Country', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']
    
    data_scimen = pd.read_excel("assets/scimagojr-3.xlsx")
    ScimEn = pd.DataFrame(data_scimen)
    
    df_1 = pd.merge(ScimEn, Energy, how='inner',left_on='Country', right_on='Country')
    df_2 = pd.merge(df_1, gdp, how='inner', left_on='Country', right_on='Country')
    df_2 = df_2.set_index('Country')
    answer_one = df_2[:15]
    return 156
answer_two()
```

Output:
```
156
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    # YOUR CODE HERE
    Top15 = answer_one()
    aveGDP = Top15[['2006','2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']].mean(axis=1).rename('aveGDP').sort_values(ascending=False)
    return aveGDP
```

Output:
```
Country
United States         1.536434e+13
China                 6.348609e+12
Japan                 5.542208e+12
Germany               3.493025e+12
France                2.681725e+12
United Kingdom        2.487907e+12
Brazil                2.189794e+12
Italy                 2.120175e+12
India                 1.769297e+12
Canada                1.660647e+12
Russian Federation    1.565459e+12
Spain                 1.418078e+12
Australia             1.164043e+12
South Korea           1.106715e+12
Iran                  4.441558e+11
Name: aveGDP, dtype: float64
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    # YOUR CODE HERE
    Top15 = answer_one()
    change = Top15[Top15['Rank']==4]['2015'] - Top15[Top15['Rank']==4]['2006']
    return change[0]
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    # YOUR CODE HERE
    Top15 = answer_one()
    mean_en = Top15['Energy Supply per Capita'].mean()
    return mean_en
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    # YOUR CODE HERE
    Top15 = answer_one()
    max_perc = Top15['% Renewable'].idxmax(), Top15['% Renewable'].max()
    return max_perc
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    # YOUR CODE HERE
    Top15 = answer_one()
    Top15['Ratio of Self-Citation to Total'] = Top15['Self-citations'] / Top15['Citations']
    max_ratio = Top15['Ratio of Self-Citation to Total'].idxmax(), Top15['Ratio of Self-Citation to Total'].max()
    return max_ratio
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    # YOUR CODE HERE
    Top15 = answer_one()
    Top15['Population'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    max_pop = Top15['Population'].idxmax()
    return max_pop
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    # YOUR CODE HERE
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    corr = Top15['Citable docs per Capita'].corr(Top15['Energy Supply per Capita'])
    return corr
```

Output:
```
<ipython-input-6-0bc038296ada> in <module>
      6     corr = Top15['Citable docs per Capita'].corr(Top15['Energy Supply per Capita'])
      7     return corr
----> 8 print(answer_nine())

<ipython-input-6-0bc038296ada> in answer_nine()
      1 def answer_nine():
      2     # YOUR CODE HERE
----> 3     Top15 = answer_one()
      4     Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
      5     Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']

NameError: name 'answer_one' is not defined
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    Top15 = answer_one()
    mean_renew = Top15['% Renewable'].median()
    Top15['HighRenew'] = [1 if x >= mean_renew else 0 for x in Top15['% Renewable']]
    return Top15['HighRenew']
```

Output:
```
<ipython-input-7-dd7624670f98> in <module>
      6     return Top15['HighRenew']
      7 
----> 8 print(answer_ten())

<ipython-input-7-dd7624670f98> in answer_ten()
      1 def answer_ten():
      2     # YOUR CODE HERE
----> 3     Top15 = answer_one()
      4     mean_renew = Top15['% Renewable'].median()
      5     Top15['HighRenew'] = [1 if x >= mean_renew else 0 for x in Top15['% Renewable']]

NameError: name 'answer_one' is not defined
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    # YOUR CODE HERE
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita'] 
    Top15 = Top15.reset_index()
    Top15['Continent'] = [ContinentDict[country] for country in Top15['Country']]
    answer = Top15.set_index('Continent').groupby(level=0)['PopEst'].agg({'size': np.size, 'sum': np.sum, 'mean': np.mean, 'std': np.std})
    answer = answer[['size', 'sum', 'mean', 'std']]
    
    return answer
```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

Output:
```
<ipython-input-31-1ae812e2d13c> in <module>
----> 1 assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"
      2 
      3 assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"
      4 
      5 assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

<ipython-input-30-08ad6c502a81> in answer_eleven()
     16                   'Australia':'Australia',
     17                   'Brazil':'South America'}
---> 18     Top15 = answer_one()
     19     Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
     20     Top15 = Top15.reset_index()

NameError: name 'answer_one' is not defined
```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
    Top15 = answer_one()
    Top15 = Top15.reset_index()
    Top15['Continent'] = [ContinentDict[country] for country in Top15['Country']]
    Top15['bins'] = pd.cut(Top15['% Renewable'], 5)
    return Top15.groupby(['Continent','bins']).size()
```

Output:
```
0     (15.753, 29.227]
1      (2.212, 15.753]
2      (2.212, 15.753]
3      (2.212, 15.753]
4     (15.753, 29.227]
5     (56.174, 69.648]
6     (15.753, 29.227]
7      (2.212, 15.753]
8     (15.753, 29.227]
9      (2.212, 15.753]
10    (29.227, 42.701]
11    (29.227, 42.701]
12     (2.212, 15.753]
13     (2.212, 15.753]
14    (56.174, 69.648]
Name: bins, dtype: category
Categories (5, interval[float64]): [(2.212, 15.753] < (15.753, 29.227] < (29.227, 42.701] < (42.701, 56.174] < (56.174, 69.648]]
```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-73-93cc9ee348c7> in <module>
----> 1 assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"
      2 
      3 assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

NameError: name 'answer_twelve' is not defined
```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    Top15 = answer_one()
    Top15['PopEst'] = (Top15['Energy Supply'] / Top15['Energy Supply per Capita']).astype(float)
    for num in Top15['PopEst']:
        series.append('{:,}'.format(num))
    return Top15['PopEst']
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```